RAS-p110α signalling in macrophages is required for effective inflammatory response and resolution of inflammation

Author response:

The following is the authors’ response to the current reviews.

Comments on revisions:

I thank the authors for addressing my comments.

- I believe that additional in vivo experiments, or the inclusion of controls for the specificity of the inhibitor, which the authors argue are beyond the scope of the current study, are essential to address the weaknesses and limitations stated in my current evaluation.

We respectfully acknowledge the reviewer's concern but would like to reiterate that demonstrating the specificity of the inhibitor is beyond the scope of this study. Alpelisib (BYL-719) is a clinically approved drug widely recognized as a specific inhibitor of p110α, primarily used in the treatment of breast cancer. Its selectivity for the p110α isoform has been extensively validated in the literature.

In our study, we used Alpelisib to assess whether pharmacological inhibition of p110α would produce effects similar to those observed in our genetic model, which is particularly relevant for the potential translational implications of our findings. Given the well-documented specificity of this inhibitor, we believe that additional controls to confirm its selectivity are unnecessary within the context of this study. Instead, our focus has been to investigate the functional role of p110α activity in macrophage-driven inflammation using the models described.

We appreciate the reviewer’s insight and hope this clarification addresses their concern.

- While the neutrophil depletion suggests neutrophils are not required for the phenotype, there are multiple other myeloid cells, in addition to macrophages, that could be contributing or accounting for the in vivo phenotype observed in the mutant strain (not macrophage specific).

We appreciate the reviewer's observation regarding the potential involvement of other myeloid cells. However, it is important to highlight that the inflammatory process follows a well-characterized sequential pattern. Our data clearly demonstrate that in the paw inflammation model:

· Neutrophils are effectively recruited, as evidenced by the inflammatory abscess filled with polymorphonuclear cells.

· However, macrophages fail to be recruited in the RBD model.

Given that this critical step is disrupted, it is reasonable to expect that any subsequent steps in the inflammatory cascade would also be affected. A precise dissection of the role of other myeloid populations would require additional lineage-specific models to selectively target each subset, which, as we have previously stated, would be the focus of an independent study.

While we cannot entirely exclude the contribution of other myeloid cells, our data strongly support the conclusion that macrophages are, at the very least, a key component of the observed phenotype. We explicitly address this point in the Discussion section, where we acknowledge the potential involvement of other myeloid populations.

- Inclusion of absolute cell numbers (in addition to the %) is essential. I do not understand why the authors are not including these data. Have they not counted the cells?

We appreciate the reviewer’s concern regarding the inclusion of absolute cell numbers. However, as stated in the Materials and Methods section, we analyzed 50,000 cells per sample, and the percentages reported in the manuscript are directly derived from this standardized count.

Our decision to present the data as percentages follows standard practices in flow cytometry-based analyses, as it allows for a clearer and more biologically relevant comparison of relative changes between conditions. This approach ensures consistency across samples and facilitates the interpretation of population dynamics during inflammation.

We would also like to clarify that all data are based on actual counts, and rigorous controls were implemented throughout the study to ensure accuracy and reproducibility. We hope this explanation addresses the reviewer’s concern and provides further clarity on our approach.

- Lastly, inclusion of representatives staining and gating strategies for all immune profiling measurements carried out by flow cytometry is important. This point has not been addressed, not even in writing.

We appreciate the reviewer’s concern regarding the inclusion of absolute cell numbers. However, as stated in the Materials and Methods section, we analyzed 50,000 cells per sample, and the percentages reported in the manuscript are directly derived from this standardized count.

Our decision to present the data as percentages follows standard practices in flow cytometry-based analyses, as it allows for a clearer and more biologically relevant comparison of relative changes between conditions. This approach ensures consistency across samples and facilitates the interpretation of population dynamics during inflammation.

We would also like to clarify that all data are based on actual counts, and rigorous controls were implemented throughout the study to ensure accuracy and reproducibility. We hope this explanation addresses the reviewer’s concern and provides further clarity on our approach.

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public review):

This study by Alejandro Rosell et al. reveals the immunoregulatory role of the RAS-p110α pathway in macrophages, specifically in regulating monocyte extravasation and lysosomal digestion during inflammation. Disrupting this pathway, through genetic tools or pharmacological intervention in mice, impairs the inflammatory response, leading to delayed resolution and more severe acute inflammation. The authors suggest that activating p110α with small molecules could be a potential therapeutic strategy for treating chronic inflammation. These findings provide important insights into the mechanisms by which p110α regulates macrophage function and the overall inflammatory response.

The updates made by the authors in the revised version have addressed the main points raised in the initial review, further improving the strength of their findings.

Reviewer #2 (Public review):

Summary:

Cell intrinsic signaling pathways controlling the function of macrophages in inflammatory processes, including in response to infection, injury or in the resolution of inflammation are incompletely understood. In this study, Rosell et al. investigate the contribution of RAS-p110α signaling to macrophage activity. p110α is a ubiquitously expressed catalytic subunit of PI3K with previously described roles in multiple biological processes including in epithelial cell growth and survival, and carcinogenesis. While previous studies have already suggested a role for RAS-p110α signaling in macrophage function, the cell intrinsic impact of disrupting the interaction between RAS and p110α in this central myeloid cell subset is not known.

Strengths:

Exploiting a sound previously described genetically engineered mouse model that allows tamoxifen-inducible disruption of the RAS-p110α pathway and using different readouts of macrophage activity in vitro and in vivo, the authors provide data consistent with their conclusion that alteration in RAS-p110α signaling impairs various but selective aspects of macrophage function in a cell-intrinsic manner.

Weaknesses:

My main concern is that for various readouts, the difference between wild-type and mutant macrophages in vitro or between wild-type and Pik3caRBD mice in vivo is modest, even if statistically significant. To further substantiate the extent of macrophage function alteration upon disruption of RAS-p110α signaling and its impact on the initiation and resolution of inflammatory responses, the manuscript would benefit from a more extensive assessment of macrophage activity and inflammatory responses in vivo.

Thank you for raising this point. We understand the reviewer’s concern regarding the modest yet statistically significant differences observed between wild-type and mutant macrophages in vitro, as well as between wild-type and Pik3caRBD mice in vivo. Our current study aimed to provide a foundational exploration of the role of RAS-p110α signaling in macrophage function and inflammatory response, focusing on a set of core readouts that demonstrate the physiological relevance of this pathway. While a more extensive in vivo assessment could offer additional insights into macrophage activity and the nuanced effects of RAS-p110α disruption, it would require an array of new experiments that are beyond the current scope.

However, we believe that the current data provide significant insights into the pathway’s role, highlighting important alterations in macrophage function and inflammatory processes due to RAS-p110α disruption. These findings lay the groundwork for future studies that can build upon our results with a more comprehensive analysis of macrophage activity in various inflammatory contexts.

In the in vivo model, all cells have disrupted RAS-p100α signaling, not only macrophages. Given that other myeloid cells besides macrophages contribute to the orchestration of inflammatory responses, it remains unclear whether the phenotype described in vivo results from impaired RAS-p100α signaling within macrophages or from defects in other haematopoietic cells such as neutrophils, dendritic cells, etc.

Thank you for raising this point. To address this, we have added a paragraph in the Discussion section acknowledging that RAS-p110α signaling disruption affects all hematopoietic cells (lines 461-470 in the discussion). However, we also provide several lines of evidence that support macrophages as the primary cell type involved in the observed phenotype. Specifically, we note that neutrophil depletion in chimera mice did not alter transendothelial extravasation, and that macrophages were the primary cell type showing significant functional defects in the paw edema model. These findings, combined with specific deficiencies in myeloid populations, suggest a predominant role of macrophages in the impaired inflammatory response, though we acknowledge the potential contributions of other myeloid cells.

Inclusion of information on the absolute number of macrophages, and total immune cells (e.g. for the spleen analysis) would help determine if the reduced frequency of macrophages represents an actual difference in their total number or rather reflects a relative decrease due to an increase in the number of other/s immune cell/s.

Thank you for this suggestion. We understand the value of presenting actual measurements; however, we opted to display normalized data to provide a clearer comparison between WT and RBD mice, as this approach highlights the relative differences in immune populations between the two groups. Normalizing data helps to focus on the specific impact of the RAS-p110α disruption by minimizing inter-sample variability that can obscure these differences.

To further address the reviewer’s concern regarding the interpretation of macrophage frequencies, we have included a pie chart that represents the relative proportions of the various immune cell populations studied within our dataset. Author response image 1 provides a visual overview of the immune cell distribution, enabling a clearer understanding of whether the observed decrease in macrophage frequency represents an actual reduction in total macrophage numbers or a shift in their relative abundance due to changes in other immune populations.

We hope this approach satisfactorily addresses reviewer’s concerns by providing both a normalized dataset for clearer interpretation of genotype-specific effects and an overall immune profile that contextualizes macrophage frequency within the broader immune cell landscape.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

As proof of concept data that activation of RAS-p110α signaling constitutes indeed a putative approach for treating chronic inflammation is not included in the manuscript, I suggest removing this implication from the abstract.

Thank you for this suggestion. We have now removed this implication from the abstract to maintain clarity and to better reflect the scope of the data presented in the manuscript.

Inclusion of a control in which RBD/- cells are also treated with BYL719, across experiments in which the inhibitor is used, would be important to determine, among other things, the specificity of the inhibitor.

We appreciate the reviewer’s suggestion to include RBD/- cells treated with BYL719 as an additional control. However, we would like to clarify that this approach would raise a different biological question, as treating RBD mice with BYL719 would not only address the specificity of the inhibitor but also examine the combined effects of genetic and pharmacologic disruptions on PI3K pathway signaling. Investigating this dual disruption falls outside the scope of our current study, which is focused specifically on the effects of RAS-p110α disruption.

It is also important to note that our RBD mouse model selectively disrupts RAS-mediated activation of p110α, while PI3K activation can still occur through other pathways, such as receptor tyrosine kinases (RTKs) and G protein-coupled receptors (GPCRs). Thus, inhibiting p110α with BYL719 would produce broader effects beyond the inhibition of RAS-PI3K signaling, impacting PI3K activation regardless of its upstream source.

In addition, incorporating this control would require us to repeat nearly all experiments in the manuscript, as it would necessitate generating and analyzing new samples for each experimental condition. Given the scope and resources involved, we believe this approach is unfeasible at this stage of the revision process.

We hope this explanation is satisfactory and that the current data in the manuscript provide a rigorous assessment of the RAS-p110α signaling pathway within the defined experimental scope.

Figure 3I is missing the statistical analysis (this is mentioned in the legend though).

Thank you for pointing this out. We apologize for the oversight. The statistical analysis for Figure 3I has now been added.

Gating strategies and representative staining should be included more generally across the manuscript.

Thank you for this suggestion. To address this, we have added a new supplementary figure (Figure 2-Supplement Figure 2) that illustrates the gating strategy along with a representative dataset. Additionally, a brief summary of the gating strategy has been included in the main text to further clarify the methodology.

It is recommended that authors show actual measurements rather than only data normalized to the control (or arbitrary units).

Thank you for this suggestion. We understand the value of presenting actual measurements; however, we opted to display normalized data to provide a clearer comparison between WT and RBD mice, as this approach highlights the relative differences in immune populations between the two groups. Normalizing data helps to focus on the specific impact of the RAS-p110α disruption by minimizing inter-sample variability that can obscure these differences.

To further address the reviewer’s concern regarding the interpretation of macrophage frequencies, we have included a pie chart that represents the relative proportions of the various immune cell populations studied within our dataset. Author response image 1 provides a visual overview of the immune cell distribution, enabling a clearer understanding of whether the observed decrease in macrophage frequency represents an actual reduction in total macrophage numbers or a shift in their relative abundance due to changes in other immune populations.

We hope this approach satisfactorily addresses reviewer’s concerns by providing both a normalized dataset for clearer interpretation of genotype-specific effects and an overall immune profile that contextualizes macrophage frequency within the broader immune cell landscape.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Macrophages are crucial in the body’s inflammatory response, with tightly regulated functions for optimal immune system performance. Our study reveals that the RAS-p110α signalling pathway, known for its involvement in various biological processes and tumorigenesis, regulates two vital aspects of the inflammatory response in macrophages: the initial monocyte movement and later-stage lysosomal function. Disrupting this pathway, either in a mouse model or through drug intervention, hampers the inflammatory response, leading to delayed resolution and the development of more severe acute inflammatory reactions in live models. This discovery uncovers a previously unknown role of the p110α isoform in immune regulation within macrophages, offering insight into the complex mechanisms governing their function during inflammation. With emerging potential to activate p110α using small molecules, targeting the RAS-p110α pathway could be a promising approach for treating chronic inflammation. This therapeutic prospect holds significant promise for easing inflammatory disorders and improving the quality of life for affected patients.

Revised version: Macrophages are crucial in the body’s inflammatory response, with tightly regulated functions for optimal immune system performance. Our study reveals that the RAS-p110α signalling pathway, known for its involvement in various biological processes and tumorigenesis, regulates two vital aspects of the inflammatory response in macrophages: the initial monocyte movement and later-stage lysosomal function. Disrupting this pathway, either in a mouse model or through drug intervention, hampers the inflammatory response, leading to delayed resolution and the development of more severe acute inflammatory reactions in live models. This discovery uncovers a previously unknown role of the p110α isoform in immune regulation within macrophages, offering insight into the complex mechanisms governing their function during inflammation and opening new avenues for modulating inflammatory responses.

Introduction

Earlier version: Phosphatidylinositol 3-kinases (PI3K) are a family of lipid kinases that phosphorylate phosphatidylinositides (PtdIns) at the 3′-hydroxyl group1. Upon activation, PI3K phosphorylates phosphatidylinositol 4,5-bisphosphate (PIP2) to generate phosphatidylinositol 3,4,5-trisphosphate (PIP3). PIP3 serves as a second messenger that recruits proteins containing pleckstrin homology (PH) domains, such as Akt (also known as protein kinase B)2, 3. This activation of PI3K regulate various cellular functions, including cell proliferation, growth, survival, motility, inflammation, and metabolism, among others1, 4. In macrophages, the activation of PI3K-Akt signalling is crucial to restrict inflammation and to promote anti-inflammatory responses in Toll Like Receptors (TLR)-induced macrophages, contributing to macrophage polarization5, 6.

Revised version: Phosphatidylinositol 3-kinases (PI3K) are a family of lipid kinases that phosphorylate phosphatidylinositides (Ptdlns) at the 3’-hydroxyl group1. Upon activation, PI3K phosphorylates phosphatidylinositol 4,5-bisphosphate (PIP2) to generate phosphatidylinositol 3,4,5-trisphosphate (PIP3). PIP3 serves as a second messenger that recruits proteins containing pleckstrin homology (PH) domains, such as Akt (also known as protein kinase B)2, 3. This activation of PI3K regulate various cellular functions, including cell proliferation, growth, survival, motility, inflammation, and metabolism, among others1, 4. In macrophages, the activation of PI3K-Akt signalling is crucial to restrict inflammation and to promote anti-inflammatory responses in Toll Like Receptors (TLR)-induced macrophages, contributing to macrophage polarization5, 6.

PI3Ks are heterodimeric lipid kinases composed of catalytic and adaptor/regulatory subunits that can be categorized into three classes based on their structures and substrate specificities3,7. Class I catalytic isoforms, including p110α, p110β, p110γ, and p110δ, play essential roles in integrating signals from growth factors, cytokines, and other environmental cues. While p110α and p110β are ubiquitously expressed, p110δ and p110γ are largely restricted to the myeloid and lymphoid lineages7–9. While p110α is primarily associated with cell growth regulation and survival in epithelial cells, it must be considered that this isoform is also expressed in immune cells, including macrophages. The role played by p110α in macrophages is not well understood, although some studies have suggested that it might regulate the survival and the regulation of the phagocytic activity of macrophages10. In the context of cancer, impairment of RAS binding to p110α in macrophages results in reduced recruitment of macrophages to the tumour site 11, 12. Additionally, this disruption leads to a change in macrophage polarization, favouring a more proinflammatory M1 state 12. These findings suggest that p110α plays a crucial role in regulating macrophage-dependent functions. However, despite these insights, the precise impact of p110α on macrophage function and the underlying molecular mechanisms influencing the inflammatory response are not yet fully understood.

Inflammation is a complex and tightly regulated series of events triggered by various stimuli such as pathogens, harmful mechanical and chemical agents, and autoimmune reactions. The inflammatory response primarily occurs in vascularized connective tissues, involving a dynamic interplay of plasma components, circulating cells, blood vessels, and cellular and extracellular factors13–17. During inflammation, mediators released by recruited leukocytes orchestrate a response that aims to facilitate tissue repair and protect the body against harmful stimuli16–18.

Macrophages play a vital role in the inflammatory response by performing functions such as antigen presentation, phagocytosis, and immunomodulation16, 17. Their role begin with the active recruitment of monocytes from the bloodstream to the site of infections19 where they differentiate into macrophages and recognize microbes and cellular debris through specific mechanisms. Macrophages subsequently actively participate in phagocytosis, a vital process involving the internalization and elimination of pathogens. Microbial destruction predominantly takes place within lysosomes and phagolysosomes20–23. At later stages of inflammation, macrophages contribute to the resolution of inflammation, thus preventing progression from acute to persistent inflammation that would cause additional tissue damage.

In this study we have used a combination of cell biology techniques and animal models to better understand the role of RAS-dependent activation of p110α at the different stages of the inflammatory response. Our findings show that RAS-p110α signalling plays a key role in the initial stages of inflammation, facilitating the extravasation of monocytes from the bloodstream by promoting the necessary cytoskeletal changes. Subsequently, RAS-p110α has a crucial role in lysosomal acidification and activation of cathepsins, which are indispensable for efficient degradation of lysosomal cargo; when these functions are impaired, prolonged acute inflammatory responses and delayed resolution steps are observed. These results significantly enhance our understanding of the complex mechanisms governing the immune response to inflammation, emphasizing the pivotal role played by RAS-p110α signalling in orchestrating proper monocyte extravasation and maintaining optimal lysosomal function.

Results

Disruption of RAS-p110α causes prolonged and more acute responses to inflammatory stress in vivo

Earlier version: Previous data suggested that, in a tumoral setting, somatic disruption of RAS-p110α prevents macrophage recruitment to tumours11, 12 and favours polarization of macrophages to a proinflammatory phenotype12, suggesting a possible role for p110α in macrophage function. To determine whether RAS-dependent activation of p110α participates in innate or adaptive immune responses to inflammation in macrophages, we used an established mouse model designed for the tamoxifen-inducible disruption of RAS binding to p110α11, 24 (Supplementary Fig. S1A). This mouse model introduces two-point mutations, T208D and K227A, in the RAS Binding Domain (RBD) of the endogenous Pik3ca allele (Pik3caRBD), enabling the selective disruption of the RAS-p110α interaction 24. Wild-type (Pik3caWT) and Pik3caRBD mice were bred with mice containing a floxed Pik3ca allele25, along with a strain containing a conditional Cre recombinase (Cre-ERT2) allele targeted to the ubiquitously expressed Rosa26 locus. The resulting Pik3caWT/Flox and Pik3caRBD/Flox mice displayed no discernible phenotype and exhibited behavior consistent with Pik3caWT mice11, 24. Activation of Cre-recombinase by tamoxifen led to the excision of the floxed Pik3ca allele, resulting in mice expressing either one Pik3caWT allele (Pik3caWT/-) or one Pik3caRBD allele (Pik3caRBD/-)11. This inducible genetic manipulation strategy allows us to selectively disrupt the RAS-p110α interaction in a controlled and temporally regulated manner, providing a valuable tool for dissecting the contributions of this pathway to immune responses in the context of inflammation.

Revised version: Previous data suggested that, in a tumoral setting, somatic disruption of RAS-p110α prevents macrophage recruitment to tumours11, 12 and favours polarization of macrophages to a proinflammatory phenotype12, suggesting a possible role for p110α in macrophage function. To determine whether RAS-dependent activation of p110α participates in innate or adaptive immune responses to inflammation in macrophages, we used an established mouse model designed for the tamoxifen-inducible disruption of RAS binding to p110α11, 24 (Figure 1-Supplement 1A). This mouse model introduces two-point mutations, T208D and K227A, in the RAS Binding Domain (RBD) of the endogenous Pik3ca allele (Pik3caRBD), enabling the selective disruption of the RAS-p110α interaction 24. Wild-type (Pik3caWT) and Pik3caRBD mice were bred with mice containing a floxed Pik3ca allele25, along with a strain containing a conditional Cre recombinase (Cre-ERT2) allele targeted to the ubiquitously expressed Rosa26 locus. The resulting Pik3caWT/Flox and Pik3caRBD/Flox mice displayed no discernible phenotype and exhibited behavior consistent with Pik3caWT mice11, 24. Activation of Cre-recombinase by tamoxifen led to the excision of the floxed Pik3ca allele, resulting in mice expressing either one Pik3caWT allele (Pik3caWT/-) or one Pik3caRBD allele (Pik3caRBD/-)11. This inducible genetic manipulation strategy allows us to selectively disrupt the RAS-p110α interaction in a controlled and temporally regulated manner, providing a valuable tool for dissecting the contributions of this pathway to immune responses in the context of inflammation.

Earlier version: Bone marrow derived macrophages (BMDMs) were generated from tibias and femurs of Pik3caWT/Flox and Pik3caRBD/Flox mice, ensuring efficient removal of the floxed allele (Supplementary Fig. S1B). Subsequently, these BMDMs were induced towards a proinflammatory state through stimulation with LPS and IFN-γ. To assess the impact of RAS-p110α binding deficiency on inflammatory intracellular signalling, we first examined Akt activation. The results revealed a decrease in Akt activation levels under inflammatory conditions in BMDMs lacking RAS-p110α binding, while no discernible change was observed in ERK activation, another well-known RAS effector (Fig. 1A and supplementary Fig. S1C and D). Interestingly, the decrease in Akt activation was accompanied by a decrease in the activation of NF-κB (Fig. 1B).

Revised version: Bone marrow derived macrophages (BMDMs) were generated from tibias and femurs of Pik3caWT/Flox and Pik3caRBD/Flox mice, ensuring efficient removal of the floxed allele (Figure 1-Supplement 1B). Subsequently, these BMDMs were induced towards a proinflammatory state through stimulation with LPS and IFN-γ. To assess the impact of RAS-p110α binding deficiency on inflammatory intracellular signalling, we first examined Akt activation. The results revealed a decrease in Akt activation levels under inflammatory conditions in BMDMs lacking RAS-p110α binding, while no discernible change was observed in ERK activation, another well-known RAS effector (Fig. 1A and Figure 1-Supplement 1C and D). Interestingly, the decrease in Akt activation was accompanied by a decrease in the activation of NF-κB (Fig. 1B).

Pik3caRBD/- mice have impaired responses to inflammatory insults.

(A) Western blotting showing activation of Akt and ERK in Pik3caWT/- and Pik3caRBD/- BMDMs activated towards a pro-inflammatory state with LPS and IFN-γ at the indicated time points; (B) Western blotting showing activation of NF-κB (p65) in Pik3caWT/- and Pik3caRBD/- BMDMs activated towards a pro-inflammatory state with LPS and IFN-γ at the indicated time points; (C) Pik3caWT/- and Pik3caRBD/- mice were injected with zymosan or PBS in the back-hind paws and inflammation was measured and plotted over time, Pik3caWT/- PBS n = 6; Pik3caRBD/- PBS n = 5; Pik3caWT/- Zymosan n = 4; Pik3caRBD/- Zymosan n = 4; Error bars indicate mean ± SEM. Significance using 2-way ANOVA test: **p < 0.01 (D) Representative H&E images of the inflamed area of Pik3caWT/- and Pik3caRBD/- paws injected with zymosan at the indicated times; (E) Representative images of cellularity present in the inflamed abscess of Pik3caWT/- and Pik3caRBD/- paws injected with zymosan at the indicated times; (F) Graph showing quantification of loose chromatin present in the inflamed abscess; (G) Representative images of macrophages (CD68 positive cells) present in the inflamed abscess of Pik3caWT/- and Pik3caRBD/- paws injected with zymosan and quantification of macrophages (CD68 positive cells) present; (H) Representative H&E images of the inflamed area and cellularity of Pik3caWT/WT paws and Pik3caWT/WT paws treated with BYL-719 injected with zymosan. (I) Representative images and quantification of macrophages (CD68 positive cells) present in the inflamed abscess of Pik3caWT/WT paws and Pik3caWT/WT paws treated with BYL-719 injected with zymosan and quantification of macrophages (CD68 positive cells). Error bars indicate mean ± SEM. Significance using Student’s t test: **p < 0.01.

Earlier version: Given the role of p65 in the expression of pro-inflammatory cytokines and chemokines26 we next analyzed the expression of various inflammation mediators using a cytokine array in unstimulated and LPS- and IFN-γ-stimulated macrophages. Results showed that, under unstimulated conditions, RAS-PI3K disruption decreases the expression of IP-10, MIP-1α (Ccl3), JE (Ccl2/MCP1), IL-16, and IL-12p70, with no upregulated cytokines observed (Supplementary Fig. S1E). The downregulation of IP-10, MIP-1α, and JE indicates an impaired ability to recruit monocytes, macrophages, and T cells to sites of inflammation. This suggests that macrophages lacking RAS-PI3K interaction may have a reduced capacity to mount a robust immune response, particularly in recruiting and activating essential immune cells needed to combat pathogens or initiate inflammation.

Revised version: Given the role of p65 in the expression of pro-inflammatory cytokines and chemokines26 we next analyzed the expression of various inflammation mediators using a cytokine array in unstimulated and LPS- and IFN-γ-stimulated macrophages. Results showed that, under unstimulated conditions, RAS-PI3K disruption decreases the expression of IP-10, MIP-1α (Ccl3), JE (Ccl2/MCP1), IL-16, and IL-12p70, with no upregulated cytokines observed (Figure 1-Supplement 1E). The downregulation of IP-10, MIP-1α, and JE indicates an impaired ability to recruit monocytes, macrophages, and T cells to sites of inflammation. This suggests that macrophages lacking RAS-PI3K interaction may have a reduced capacity to mount a robust immune response, particularly in recruiting and activating essential immune cells needed to combat pathogens or initiate inflammation.

Earlier version: Regarding LP +IFN-γ stimulated macrophages, results showed a decrease in the expression of IL-1β and IL-17, key drivers of pro-inflammatory responses, and an upregulation of IL-7, ILra, JE, BLC, I-309, Eotaxin, and G-CSF (Supplementary Fig. S1F). Elevated levels of these factors are associated with enhanced chemotactic signals and regulatory functions. Thus, the cytokine and chemokine expression profile observed in RAS-PI3K deficient macrophages suggests impaired pro-inflammatory responses and altered immune cell recruitment patterns, potentially influencing the resolution of inflammation and tissue repair processes.

Revised version: Regarding LP +IFN-γ stimulated macrophages, results showed a decrease in the expression of IL-1β and IL-17, key drivers of pro-inflammatory responses, and an upregulation of IL-7, Ilra, JE, BLC, I-309, Eotaxin, and G-CSF (Figure 1-Supplement 1F). Elevated levels of these factors are associated with enhanced chemotactic signals and regulatory functions. Thus, the cytokine and chemokine expression profile observed in RAS-PI3K deficient macrophages suggests impaired pro-inflammatory responses and altered immune cell recruitment patterns, potentially influencing the resolution of inflammation and tissue repair processes.

Earlier version: Next we assessed whether the absence of RAS binding to p110α affects the ability of BMDMs to acquire a pro-inflammatory state characterized by increased expression of markers such as CD80, CD86, and MHCII27, 28. To assess this, Pik3caRBD/-and Pik3caWT/- BMDMs were stimulated with LPS+IFN-γ, followed by flow cytometry analysis. Expression levels of CD80 (Supplementary Fig. S1G), CD86 (Supplementary Fig. S1H), and MCHII (Supplementary Fig. S1I) were examined. No significant differences were observed in the expression levels of any of these markers between the two genotypes under study.

Revised version: Next we assessed whether the absence of RAS binding to p110α affects the ability of BMDMs to acquire a pro-inflammatory state characterized by increased expression of markers such as CD80, CD86, and MHCII27, 28. To assess this, Pik3caRBD/-and Pik3caWT/- BMDMs were stimulated with LPS+IFN-γ, followed by flow cytometry analysis. Expression levels of CD80 (Figure 1-Supplement 1G), CD86 (Figure 1-Supplement 1H), and MCHII (Figure 1-Supplement 1I) were examined. No significant differences were observed in the expression levels of any of these markers between the two genotypes under study.

Earlier version: Consequently, our next objective was to investigate whether in vivo disruption of RAS-p110α might lead to altered inflammatory responses. To address this, we administered tamoxifen to 10–12-week-old Pik3caWT/flox and Pik3caRBD/flox mice with tamoxifen and, after a two-week interval, conducted a paw swelling assay. In this assay, zymosan (10 μg/μl) or PBS was injected into the hind paws of Pik3caRBD/- and Pik3caWT/- mice with paw thickness measured at regular intervals over a 5-day period. The results revealed a significant increase in paw inflammation in Pik3caRBD/- mice compared to Pik3caWT/- mice, evident from the earlier time points measured and persisting throughout the experiment (Fig. 1D and 1E). Paws from the Pik3caWT/flox and Pik3caRBD/flox mice, which had not received tamoxifen, exhibited comparable levels of inflammation (Supplementary Fig. S1J). This reaffirms that the absence of p110α triggers a significant alteration in the inflammatory response and confirms that Pik3caWT/flox and Pik3caRBD/flox mice do not show any phenotype, as previously described11, 24. Additionally, analysis of the blood sedimentation rate, a reliable indicator of systemic inflammation levels29, showed a lower basal sedimentation ratio in Pik3caRBD/- mice under PBS-treated conditions (Supplementary Fig. S1K). Following zymosan injection, both genotypes exhibited an increase in sedimentation rate, but the rise was more pronounced in Pik3caRBD/- mice, indicating a heightened systemic inflammatory response. These findings underscore that disruption of RAS-p110α interaction results in an exacerbated inflammatory state, reflected in both localized paw inflammation and systemic inflammatory mediator levels.

Revised version: Consequently, our next objective was to investigate whether in vivo disruption of RAS-p110α might lead to altered inflammatory responses. To address this, we administered tamoxifen to 10-12-week-old Pik3caWT/flox and Pik3caRBD/flox mice with tamoxifen and, after a two-week interval, conducted a paw swelling assay. In this assay, zymosan (10 μg/μl) or PBS was injected into the hind paws of Pik3caRBD/- and Pik3caWT/- mice with paw thickness measured at regular intervals over a 5-day period. The results revealed a significant increase in paw inflammation in Pik3caRBD/- mice compared to Pik3caWT/- mice, evident from the earlier time points measured and persisting throughout the experiment (Fig. 1D and 1E). Paws from the Pik3caWT/flox and Pik3caRBD/flox mice, which had not received tamoxifen, exhibited comparable levels of inflammation (Figure 1-Supplement 1J). This reaffirms that the absence of p110α triggers a significant alteration in the inflammatory response and confirms that Pik3caWT/flox and Pik3caRBD/flox mice do not show any phenotype, as previously described11, 24. Additionally, analysis of the blood sedimentation rate, a reliable indicator of systemic inflammation levels29, showed a lower basal sedimentation ratio in Pik3caRBD/- mice under PBS-treated conditions (Figure 1-Supplement 1K). Following zymosan injection, both genotypes exhibited an increase in sedimentation rate, but the rise was more pronounced in Pik3caRBD/- mice, indicating a heightened systemic inflammatory response. These findings underscore that disruption of RAS-p110α interaction results in an exacerbated inflammatory state, reflected in both localized paw inflammation and systemic inflammatory mediator levels.

To delve deeper into the inflammatory response induced by zymosan, we collected paw samples at two- and five-days post-injection and conducted Haematoxylin & Eosin (H&E) studies. This approach enabled a comprehensive analysis, allowing us not only to scrutinize the early stages of inflammation but also to monitor the subsequent resolution phase of the inflammatory process. Paws obtained from mice injected with zymosan for 2 days displayed extensive areas of damaged connective tissue in both Pik3caRBD/- and Pik3caWT/- paws (Fig. 1D). Notably, the inflamed region in Pik3caRBD/- mice was larger compared to control samples. After 5 days of zymosan injection, there was a significant reduction in the inflamed area, although it remained comparatively larger in the paws from Pik3caRBD/- mice (Fig. 1D). This observation suggests a prolonged and heightened inflammatory response in mice lacking RAS-p110α interaction, emphasizing the role of this interaction in the regulation of inflammatory processes.

The cellular composition within the inflammatory abscess offers crucial insights into the severity and progression of the disease. Close examination with the pathologist revealed features indicative of an acute inflammatory response, including an inflammatory abscess with elevated numbers of polymorphonuclear cells, primarily neutrophils, and macrophages displaying altered cell shape and increased cell death, accompanied by fibrin and chromatin deposition (Fig. 1E and F). Notably, Pik3caRBD/- mice exhibited lower numbers of macrophages, a larger necrotic area with increased chromatin remnants, and reduced fibrin content compared to the paws from Pik3caWT/- mice (Fig. 1E and F). By day 5, paws from Pik3caWT/- mice showed a significant increase in the number of macrophages, a decrease in polymorphonuclear cells, and a nearly complete resolution of the necrotic area, indicating the initiation of the resolution phase. Moreover, there were abundant activated fibroblasts), suggesting active production of new connective tissue. In contrast, paws from Pik3caRBD/- mice exhibited a delayed healing process characterized by a larger central area of polymorphonuclear cells, abundant fibrin deposits, reduced numbers of infiltrating macrophages, and limited fibroblast activity (Fig. 1E and F). These findings collectively indicate an imbalance in the inflammatory response and a slower progression towards resolution in the absence of RAS-p110α interaction, emphasizing the pivotal role of this interaction in orchestrating an effective and timely resolution of inflammation.

To evaluate the presence of macrophages within the inflammatory lesion, we performed specific immunohistochemical (IHC) analysis using the macrophage-specific marker CD68 in paws from day 5, where more cellular preservation and initiation of the healing process had been observed. A notable decrease in the number of CD68-positive cells was observed in the inflamed abscess region of Pik3caRBD/- mice (Fig. 1G).

Earlier version: To further confirm the involvement of p110α signalling in the acute inflammatory response, Pik3caWT/WT mice were subjected to daily treatment with BYL719 (Alpelisib), a specific inhibitor of p110α isoform. After an initial 48-hour treatment period, the mice received injections of zymosan or PBS into their back-hind paws and were sacrificed 2 days later for analysis. Similar to what we had observed in Pik3caRBD/-mice, the inflamed area in BYL719-treated mice exhibited a larger extension than the inflamed are from non-treated mice (Fig. 1H and Supplementary Fig. S1L). The central necrotic region was significantly larger in the BYL719-treated mice, and it contained large amount of apoptotic polymorphonuclear cells, lower number of macrophages, and increased deposits of chromatin (Fig. 1H) resembling the observations from Pik3caRBD/- mice. Immunostaining for CD68 revealed a reduction in the number of macrophages present in the inflammatory abscess upon inhibition of p110α signalling with BYL719 treatment (Fig. 1I).

Revised version: To further confirm the involvement of p110α signalling in the acute inflammatory response, Pik3caWT/WT mice were subjected to daily treatment with BYL719 (Alpelisib), a specific inhibitor of p110α isoform. After an initial 48-hour treatment period, the mice received injections of zymosan or PBS into their back-hind paws and were sacrificed 2 days later for analysis. Similar to what we had observed in Pik3caRBD/-mice, the inflamed area in BYL719-treated mice exhibited a larger extension than the inflamed are from non-treated mice (Fig. 1H and Figure 1-Supplement 1L). The central necrotic region was significantly larger in the BYL719-treated mice, and it contained large amount of apoptotic polymorphonuclear cells, lower number of macrophages, and increased deposits of chromatin (Fig. 1H) resembling the observations from Pik3caRBD/- mice. Immunostaining for CD68 revealed a reduction in the number of macrophages present in the inflammatory abscess upon inhibition of p110α signalling with BYL719 treatment (Fig. 1I).

In summary, our findings highlight that both genetic disruption of RAS-p110α interaction and pharmacological inhibition of p110α contribute to an expanded inflamed area and central necrotic region, concurrently reducing macrophage infiltration in a zymosan-induced inflammation model. These results collectively underscore the pivotal role of the p110α isoform of PI3K in orchestrating the resolution of inflammatory responses, emphasizing its significance in modulating the dynamics and outcomes of inflammatory processes.

Disruption of RAS binding to p110α impairs the number of inflammatory monocytes in blood and spleen

Earlier version: Given the decrease in macrophages observed in the inflammatory abscess, we next set out to determine whether disruption of RAS binding to p110α had an effect on the number of monocytes circulating in the blood of adult mice. Pik3caWT/flox and Pik3caRBD/flox mice were treated with tamoxifen at 12–14 weeks of age and 4 weeks later, blood was collected by cardiac puncture and immune populations were analysed by flow cytometry. We found a decrease in the number of circulating classical (or inflammatory) monocytes (Ly6CHi/Ly6G–/CD11b+ cells) in Pik3caRBD/- mice (Fig. 2A) and no changes were observed in non-classical (or non-inflammatory) monocytes (Ly6CLo/Ly6G–/CD11b+ cells) (Fig. 2B). Together with the decrease in inflammatory monocytes, we observed an increase in the number of neutrophils (Ly6C–/Ly6G+/CD11b+ cells) (Supplementary Fig. S2A). We did not detect differences in the numbers of T-cells (CD3+, CD8+ or CD4+) (Supplementary Fig. S2B) or B-cells (CD19+) (Supplementary Fig. S2C) in the blood after the disruption of RAS binding to p110α. We also analysed these same cell populations in Pik3caWT/flox and Pik3caRBD/flox mice and no differences were found (data not shown).

Revised version: Given the decrease in macrophages observed in the inflammatory abscess, we next set out to determine whether disruption of RAS binding to p110α had an effect on the number of monocytes circulating in the blood of adult mice. Pik3caWT/flox and Pik3caRBD/flox mice were treated with tamoxifen at 12-14 weeks of age and 4 weeks later, blood was collected by cardiac puncture and immune populations were analysed by flow cytometry. To assess the effects of RAS-p110α disruption on immune populations, we employed a flow cytometry gating strategy as described in the Methods section. Briefly, CD45+ cells were initially gated to define the overall immune cell population. Within this population, B cells were identified by CD19+ CD11b-staining, and T cells were defined by CD3+ expression. Further differentiation of T cell subtypes (CD4+, CD8+, and double-negative T cells) was conducted based on CD4 and CD8 markers within the CD3+ cell population. Additionally, myeloid cells were gated as CD11b+ cells, with further identification of granulocytes and monocytes based on Ly6G and Ly6C expression. Specifically, granulocytes/neutrophils were characterized as Ly6G+ Ly6C+, classical monocytes as Ly6G- Ly6C+, and alternative monocytes as Ly6G- Ly6C- (Figure 2-Supplement 2A). We found a decrease in the number of circulating classical (or inflammatory) monocytes (Ly6CHi/Ly6G–/CD11b+ cells) in Pik3caRBD/- mice (Fig. 2A) and no changes were observed in non-classical (or non-inflammatory) monocytes (Ly6CLo/Ly6G–/CD11b+ cells) (Fig. 2B). Together with the decrease in inflammatory monocytes, we observed an increase in the number of neutrophils (Ly6C–/Ly6G+/CD11b+ cells) (Figure 2-Supplement 3A). We did not detect differences in the numbers of T-cells (CD3+, CD8+ or CD4+) (Figure 2-Supplement 3B) or B-cells (CD19+) (Figure 2-Supplement 3C) in the blood after the disruption of RAS binding to p110α. We also analysed these same cell populations in Pik3caWT/flox and Pik3caRBD/flox mice and no differences were found (data not shown).

Earlier version: Disruption of RAS-p110α interaction decreases the number of inflammatory monocytes in blood and spleen.

Revised version: Disruption of RAS-p110a interaction decreases the number of inflammatory monocytes in blood and spleen.

Twelve-week-old mice were treated with tamoxifen and after four weeks, flow cytometry analysis was performed to determine (A) inflammatory monocytes in circulating blood; (B) alternatively activated monocytes in circulating blood; (C) inflammatory monocytes in spleen; (D) classically activated monocytes in spleen; (E) macrophages in spleen; (F) macrophages in spleen’s white pulp; (G) macrophages in spleen’s red pulp. Data are presented as percentage of positive cells for the indicated markers. Black dots represent data from Pik3caWT/- mice and red dots represent data from Pik3caRBD/- mice. Each dot represents an individual mouse. Error bars indicate mean ± SEM. Significance using Student’s t test: * p < 0.05; **p < 0.01.

Earlier version: Since splenic monocytes resemble their blood counterparts30 we next aimed at determining whether splenic monocyte population would also be altered after disruption of RAS p110α interaction. As observed in circulating blood, the number of classical monocytes (Ly6CHi/Ly6G–/CD11b+) in the spleen decreased after disruption of RAS-p110α interaction (Fig. 2C) and no differences were found in non-classical activated monocytes (Ly6CLo/ Ly6G–/CD11b+) (Fig. 2D). We also checked the levels of resident macrophages present in the spleen and results showed that spleens from the Pik3caRBD/- mice had a decrease in the number of differentiated macrophages (F4/80+/CD11b+/CD45+) (Fig. 2E). Analysis of macrophages in the white and red pulp of the spleen indicated that Pik3caRBD/- mice had a significant decrease in the macrophage population in the former region (Fig. 2F and 2G). No differences were found in the number of granulocytes (Supplementary Fig. S2D), B-cells (Supplementary Fig. S2E) or T-cells (Supplementary Fig. S2F) present in the spleen between Pik3caRBD/- and Pik3caWT/- mice. There were no differences in spleen size from Pik3caRBD/- and Pik3caWT/- mice either (Supplementary Fig. S2G).

Revised version: Since splenic monocytes resemble their blood counterparts30 we next aimed at determining whether splenic monocyte population would also be altered after disruption of RAS p110α interaction. Flow cytometry from Pik3caWT/- and Pik3caRBD/- mice was performed following the same gating strategy as described for blood samples (Figure 2-Supplement 2B). As observed in circulating blood, the number of classical monocytes (Ly6CHi/Ly6G–/CD11b+) in the spleen decreased after disruption of RAS-p110α interaction (Fig. 2C) and no differences were found in non-classical activated monocytes (Ly6CLo/ Ly6G– /CD11b+) (Fig. 2D). We also checked the levels of resident macrophages present in the spleen and results showed that spleens from the Pik3caRBD/- mice had a decrease in the number of differentiated macrophages (F4/80+/CD11b+/CD45+) (Fig. 2E). Analysis of macrophages in the white and red pulp of the spleen indicated that Pik3caRBD/- mice had a significant decrease in the macrophage population in the former region (Fig. 2F and 2G). No differences were found in the number of granulocytes (Figure 2-Supplement 3D), B-cells (Figure 2-Supplement 3E) or T-cells (Figure 2-Supplement 3F) present in the spleen between Pik3caRBD/- and Pik3caWT/- mice. There were no differences in spleen size from Pik3caRBD/- and Pik3caWT/- mice either (Figure 2-Supplement 3G).

Earlier version: Given the decrease in the number of inflammatory monocytes and macrophages observed in blood and spleens of Pik3caRBD/- mice, we wondered if disruption of RAS activation of p110α could lead to a decrease in myeloid bone marrow precursors. Myeloid lineage descends from a common myeloid progenitor (CMP) in bone marrow and traverse into blood as mature cells31. CMP differentiates into the Granulocyte-Macrophage Progenitor (GMP) and the Megakaryocyte-Erythroid Progenitor (MEP). The CMP can generate all types of myeloid colonies, whereas the GMP or the MEP produces only granulocyte macrophage (GM) or megakaryocyte erythrocyte (MegE) lineage cells, respectively (Supplementary Fig. S2H). No significant differences were noted in the quantity of CMP progenitors within the bone marrow Pik3caRBD/- and Pik3caWT/- mice (Supplementary Fig. S2I-L). This observation suggests that the variations observed in blood and spleen parameters are not attributable to impairments in progenitor fate.

Revised version: Given the decrease in the number of inflammatory monocytes and macrophages observed in blood and spleens of Pik3caRBD/- mice, we wondered if disruption of RAS activation of p110α could lead to a decrease in myeloid bone marrow precursors. Myeloid lineage descends from a common myeloid progenitor (CMP) in bone marrow and traverse into blood as mature cells31. CMP differentiates into the Granulocyte-Macrophage Progenitor (GMP) and the Megakaryocyte-Erythroid Progenitor (MEP). The CMP can generate all types of myeloid colonies, whereas the GMP or the MEP produces only granulocyte macrophage (GM) or megakaryocyte erythrocyte (MegE) lineage cells, respectively (Figure 2-Supplement 3H). No significant differences were noted in the quantity of CMP progenitors within the bone marrow Pik3caRBD/- and Pik3caWT/- mice (Figure 2-Supplement 3I-L). This observation suggests that the variations observed in blood and spleen parameters are not attributable to impairments in progenitor fate.

Earlier version: Finally, bone marrow precursors from Pik3caRBD/flox and Pik3caWT/flox mice were differentiated into macrophages in vitro, and the number of bone marrow-derived macrophages (BMDM) was determined by flow cytometry analysis. No differences were found in the number of BMDMs obtained from Pik3caRBD/- and Pik3caWT/- mice (Supplementary Fig. S2M), suggesting that the disruption of RAS-p110α signalling do not interfere with the ability of bone marrow precursors to differentiate to macrophages.

Revised version: Finally, bone marrow precursors from Pik3caRBD/flox and Pik3caWT/flox mice were differentiated into macrophages in vitro, and the number of bone marrow-derived macrophages (BMDM) was determined by flow cytometry analysis. No differences were found in the number of BMDMs obtained from Pik3caRBD/- and Pik3caWT/- mice (Figure 2-Supplement 3M), suggesting that the disruption of RAS-p110a signalling do not interfere with the ability of bone marrow precursors to differentiate to macrophages.

Disruption of RAS binding to p110α impairs monocyte transendothelial extravasation in response to inflammatory cues

The decrease in classical monocytes observed in the inflammatory abscess of paws from Pik3caRBD/- mice may indicate a decreased ability to mount an effective immune response. We carried out transwell assays since they are widely used to quantify transendothelial migration. Fibroblasts were seeded in the lower chamber to provide a continuous supply of Ccl232, as this cytokine is well known for its ability to drive chemotaxis of myeloid cells under inflammatory conditions33. Pik3caRBD/- or Pik3caWT/- BMDMs were seeded in the upper chamber of the transwell. Additionally, we stimulated Pik3caRBD/- or Pik3caWT/- BMDMs with lipopolysaccharide (LPS) and Interferon gamma (IFN-γ) to mimic/recapitulate the proinflammatory phenotype typically associated with bacterial infection that causes macrophage activation34. Results showed that in both unstimulated and LPS+IFN-γ stimulated BMDMs, a lower number of Pik3caRBD/- BMDMs were able to go through the trans-well pore membranes compared to Pik3caWT/- BMDMs (Fig. 3A). As expected, when BMDMs were stimulated towards a proinflammatory phenotype, a decrease in their migratory ability is observed35.

Earlier version: Disruption of RAS-p110α interaction impairs transendothelial extravasation to sites of inflammation.

Revised version: Disruption of RAS-p110a interaction impairs transendothelial extravasation to sites of inflammation.

Earlier version: (A) Graph indicating the quantification of Pik3caRBD/- and Pik3caRBD/- BMDMs passing through 8 μm membrane pore transwells for 24 hours. Membrane attached macrophages on the lower part of the transwell were stained with crystal violet and quantified. Each dot represents an independent experiment. Error bars indicate mean ± SEM; (B) Random migration of Pik3caWT/-, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 (500 ng/ml) was analysed by time-lapse video microscopy and cell tracing in the presence or absence of LPS (100 ng/ml) and IFN-γ (20 ng/ml); (C) Schematic representation of myeloid chimera generation strategy; (D) Graph showing level of bone marrow reconstitution with Pik3caRBD/- and Pik3caWT/- myeloid lineage; (E) Graph quantifying the number of neutrophils in the blood of chimera mice treated with anti-GRl (25μg/mouse/day); (F-H) Intra vital imaging quantification of the number of: (F) rolling monocytes per 5 minutes; (G) adherent monocytes per 500 μm of vessel segment; (H) extravasated monocytes per mm2 of tissue. Each dot represents an individual mouse. (I) Chimera Pik3caWT/- and Pik3caRBD/- mice were injected with zymosan or PBS in the back-hind paws and inflammation was measured and plotted over time, Pik3caWT/- PBS n = 5; Pik3ca’RBD/- PBS n = 5; Pik3caWT/-Zymosan n = 5; Pik3caRBD/- Zymosan n = 5.

Revised version: (A) Graph indicating the quantification of Pik3caRBD/- and Pik3caRBD/- BMDMs passing through 8 μm membrane pore transwells for 24 hours. Membrane attached macrophages on the lower part of the transwell were stained with crystal violet and quantified. Each dot represents an independent experiment. Error bars indicate mean ± SEM; (B) Random migration of Pik3caWT/-, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 (500 ng/ml) was analysed by time-lapse video microscopy and cell tracing in the presence or absence of LPS (100 ng/ml) and IFN-γ (20 ng/ml); (C) Schematic representation of myeloid chimera generation strategy; (D) Graph showing level of bone marrow reconstitution with Pik3caRBD/- and Pik3caWT/- myeloid lineage; (E) Graph quantifying the number of neutrophils in the blood of chimera mice treated with anti-GR1 (25μg/mouse/day); (F-H) Intra vital imaging quantification of the number of: (F) rolling monocytes per 5 minutes; (G) adherent monocytes per 500 μm of vessel segment; (H) extravasated monocytes per mm2 of tissue. Each dot represents an individual mouse. (I) Chimera Pik3caWT/- and Pik3caRBD/- mice were injected with zymosan or PBS in the back-hind paws and inflammation was measured and plotted over time, Pik3caWT/- PBS n = 5; Pik3caRBD/- PBS n = 5; Pik3caWT/-Zymosan n = 5; Pik3caRBD/- Zymosan n = 5.

Earlier version: Statistical significance was obtained using Mann-Whitney test (A-H) or 2-way ANOVA test (I): n.s. nonsignificant; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001.

Revised version: Statistical significance was obtained using Mann-Whitney test (A-H) or 2-way ANOVA test (I): n.s. non significant; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001.

Extravasation entails the migration of monocytes through the endothelium36, 37, so we conducted random migration assays of Pik3caRBD/- and Pik3caWT/- BMDMs growing in matrigel coated plates under unstimulated conditions or activated towards an inflammatory phenotype by addition of LPS and IFN-γ to the media. Analysis of the data revealed that, under pro-inflammatory conditions, Pik3caRBD/- BMDMs displayed a decrease in migration speed (Fig. 3B). Migration was also evaluated in Pik3caWT/WT BMDMs treated with BYL-719. Results confirmed a decrease in migration speed of macrophages upon treatment with BYL-719 (Fig. 3B) indicating that inhibition of p110α, either genetically or pharmacologically, reduces the ability of macrophages to migrate under inflammatory conditions.

Earlier version: To determine if disruption of RAS binding to p110α impairs monocyte ability to extravasate through the endothelium in vivo in response to an inflammatory stress, we analysed monocyte extravasation through the mesenteric vein in response to intraperitoneal Ccl2 injection. It is well established that loss of p110α function leads to significant impairment in endothelial and lymphatic system12, 38, so in order to determine if monocytes from Pik3caRBD/- mice presented an alteration in extravasation, we generated chimeric mice in which only the bone marrow were defective in RAS binding to p110α (Fig. 3C). For this, bone marrow from Pik3caRBD/- or Pik3caWT/- mice was injected through the tail vein of irradiated LysM-GFP donor mice39. Engraftment of Pik3caRBD/- and Pik3caWT/- bone marrow could be followed by disappearance of the eGFP signal in donor mice (Fig. 3D and Supplementary Fig. S3A). Neutrophils (but no monocytes) were depleted using an anti-GR1 antibody (Fig. 3E and Supplementary Fig. S3B) to avoid interference with monocyte extravasation. Intraperitoneal injection of Ccl2 was performed in the peritoneum of the chimera mice to induce extravasation of monocytes through the mesenteric vein and rolling, adhesion and extravasation was measured by intravital microscopy. Flow cytometry analysis demonstrated no differences in Ccr2 expression between Pik3caRBD/- and Pik3caWT/- monocytes (Supplementary Fig. S3C). We observed that blocking RAS-p110α interaction did not induce a decrease in the number of rolling monocytes (Fig. 3F) or in the number of monocytes that adhere to the endothelium (Fig. 3G). However, when extravasation was measured, data showed that disruption of RAS binding to p110α caused a significant decrease in the number of monocytes that were capable of extravasating through the endothelium of the mesenteric vein (Fig. 3H).

Revised version: To determine if disruption of RAS binding to p110α impairs monocyte ability to extravasate through the endothelium in vivo in response to an inflammatory stress, we analysed monocyte extravasation through the mesenteric vein in response to intraperitoneal Ccl2 injection. It is well established that loss of p110α function leads to significant impairment in endothelial and lymphatic system12, 38, so in order to determine if monocytes from Pik3caRBD/- mice presented an alteration in extravasation, we generated chimeric mice in which only the bone marrow were defective in RAS binding to p110α (Fig. 3C). For this, bone marrow from Pik3caRBD/- or Pik3caWT/- mice was injected through the tail vein of irradiated LysM-GFP donor mice39. Engraftment of Pik3caRBD/- and Pik3caWT/- bone marrow could be followed by disappearance of the eGFP signal in donor mice (Fig. 3D and Figure 3-Supplement 4A).

Added in the revised version: Neutrophils (but no monocytes) were depleted using an anti-GR1 antibody (Fig. 3E and Figure 3- Supplement 4B) to avoid interference with monocyte extravasation. Intraperitoneal injection of Ccl2 was performed in the peritoneum of the chimera mice to induce extravasation of monocytes through the mesenteric vein and rolling, adhesion and extravasation was measured by intravital microscopy. Flow cytometry analysis demonstrated no differences in Ccr2 expression between Pik3caRBD/- and Pik3caWT/- monocytes (Figure 3-Supplement 4C). We observed that blocking RAS-p110α interaction did not induce a decrease in the number of rolling monocytes (Fig. 3F) or in the number of monocytes that adhere to the endothelium (Fig. 3G). However, when extravasation was measured, data showed that disruption of RAS binding to p110α caused a significant decrease in the number of monocytes that were capable of extravasating through the endothelium of the mesenteric vein (Fig. 3H).

Additionally, in order to determine whether the differences observed in the inflammatory response of Pik3caRBD/- mice were due to lack of monocyte extravasation, we repeated the zymosan paw swelling assay in the chimera mice. As previously, we observed that disruption of RAS-p110α only in the immune system led to higher inflammatory response and a delay in the initiation of the resolution phase (Fig. 3I).

Disruption of RAS-p110α activation in macrophages induces changes in cytoskeleton reorganization

During transendothelial migration, leukocytes undergo cytoskeletal rearrangements that allow them to squeeze through the tight spaces between endothelial cells and enter the underlying tissue40, 41. Therefore, we next sought to determine whether the differences observed in Pik3caRBD/- BMDMs extravasation and migration were attributable to altered cytoskeletal dynamics. To do so, we first aimed at examining cell shape and spread in Pik3caRBD/- and Pik3caWT/- BMDMs in pro-inflammatory conditions after LPS+IFN-γ treatment. Treatment with LPS+IFNγ induced an increase in cell spread in both Pik3caRBD/- and Pik3caWT/- BMDMs when compared to their respective unstimulated counterparts (Fig. 4A and 4B). However, cell spread in LPS+IFNγ activated Pik3caRBD/- BMDMs was significantly decreased compared to that observed Pik3caWT/- BMDMs (Fig. 4A and 4B). Additionally, Pik3caRBD/- BMDMs were more elongated and did not acquire the typical rounded shape known to be induced in macrophages after treatment with LPS+IFN-γ42 (Fig. 4A and 4C). We also analysed cell height and results showed that Pik3caRBD/- BMDMs are higher both in unstimulated conditions and after LPS+IFN-γ stimulation (Fig. 4D).

Earlier version: Disruption of RAS-p110α signalling impairs BMDM ability to remodel their cytoskeleton in response to inflammatory cues.

Revised version: Disruption of RAS-p110a signalling impairs BMDM ability to remodel their cytoskeleton in response to inflammatory cues.

(A) Representative images of Pik3caWT/- and Pik3caRBD/- BMDMs unstimulated or activated with LPS+ IFN-γ; (B) violin plot quantifying spread area. IF images of Pik3caWT/- and Pik3caRBD/- BMDMs co-stained with phalloidin and DAPI were used to analyse spread area. Results are represented as a violin plot. Three independent biological replicates were analysed (n ≥ 250 total cells); (C) Quantification of cell circularity of Pik3caWT/- and Pik3caRBD/- BMDMs unstimulated or activated with LPS+ IFN-γ. Quantification was performed using the same images from panel (A); (D) Representative 3D projections of Pik3caWT/- and Pik3caRBD/- BMDMs unstimulated or activated with LPS+IFN-γ and violin plot showing quantification of the corresponding cell height. 3D projections and cell height were analysed using same images used in panel (A). (E) Violin plot representing F-actin pool in Pik3caWT/- and Pik3caRBD/- BMDMs unstimulated or activated with LPS+ IFN-γ. Three independent biological replicates were analysed (n ࣙ 250 total cells); (F) Violin plot representing G-actin pool (measured by DNAse staining) in Pik3caWT/- and Pik3caRBD/- BMDMs unstimulated or activated with LPS+ IFN-γ. Three independent biological replicates were analysed (n ≥ 250 total cells); (G) Graph showing stiffness (elastic modulus) of Pik3caWT/- and Pik3caRBD/- BMDMs; (H) Graph showing phagocytosis of apoptotic cells (efferocytosis) over time in Pik3caWT/-, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 BMDMs.

Statistical significance was obtained using Mann-Whitney test: n.s. non-significant; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001.

Earlier version: Actin is a bona fide regulator of cell shape43, so we next aimed at exploring the actin cytoskeleton in Pik3caRBD/- and Pik3caWT/- BMDMs. Actin fractionation assays were carried out in unstimulated or LPS+IFNY activated Pik3caRBD/- and Pik3caWT/- BMDMs. Our results revealed that disruption of RAS-p110α binding caused a decrease in the F-actin pool in unstimulated BMDMs (Fig. 4E), with no changes observed in the G-actin pool when compared to matched controls (Fig. 4F). Analysis of actin dynamics after activation with LPS+IFNγ showed that, in proinflammatory conditions, the F-actin pool is increased, as expected44 (Fig. 4E). However, actin polymerization was significantly active in Pik3caRBD/- BMDMs, leading to a striking increase in F-actin when compared to that observed in controls (Fig. 4E). In parallel, a decrease in the G-actin pool in Pik3caRBD/- BMDMs was observed, suggesting an increase in the stabilization of F-actin after disruption of RAS binding to p110α (Fig. 4F).

Revised version: Actin is a bona fide regulator of cell shape43, so we next aimed at exploring the actin cytoskeleton in Pik3caRBD/- and Pik3caWT/- BMDMs. Actin fractionation assays were carried out in unstimulated or LPS+IFNγ activated Pik3caRBD/- and Pik3caWT/- BMDMs. Our results revealed that disruption of RAS-p110α binding caused a decrease in the F-actin pool in unstimulated BMDMs (Fig. 4E), with no changes observed in the G-actin pool when compared to matched controls (Fig. 4F). Analysis of actin dynamics after activation with LPS+IFNγ showed that, in proinflammatory conditions, the F-actin pool is increased, as expected44 (Fig. 4E). However, actin polymerization was significantly active in Pik3caRBD/- BMDMs, leading to a striking increase in F-actin when compared to that observed in controls (Fig. 4E). In parallel, a decrease in the G-actin pool in Pik3caRBD/- BMDMs was observed, suggesting an increase in the stabilization of F-actin after disruption of RAS binding to p110α (Fig. 4F).

We next sought to investigate whether the enhanced F-actin stabilization corresponded to increased cellular rigidity. Consequently, we measured the elastic modulus of Pik3caRBD/- BMDMs, revealing a significant increase in cell stiffness following RAS-p110α disruption under both basal and proinflammatory conditions (Fig. 4G). To validate that the heightened cell stiffness was attributed to p110α, we treated Pik3caWT/WT BMDMs with BYL-719 and observed an amplified elastic modulus in these cells when stimulated with LPS+IFN-γ (Fig. 4G), confirming that the inhibition of p110α indeed results in augmented cellular rigidity.

Earlier version: These observations led us to hypothesize that phagocytosis, which heavily relies on actin rearrangement, might also be affected in Pik3caRBD/- BMDMs. To test this hypothesis, we assessed the ability of Pik3caRBD/- and Pik3caWT/- BMDMs to phagocytose fluorescent microspheres, Borrelia burgdorferi, and apoptotic cells. The results revealed varying outcomes depending on the target. When phagocytosing 1 μm non-opsonized green-fluorescent beads, Pik3caRBD/- BMDMs were less efficient at engulfing the microspheres compared to their WT counterparts (Supplementary Fig. S4A). Similar results were obtained when Pik3caWT/WT BMDMs were treated with BYL-719 (Supplementary Fig. S4A). In contrast, no significant differences were observed between Pik3caRBD/- and Pik3caWT/- BMDMs in their ability to phagocytose Borrelia burgdorferi (Supplementary Fig. S4B).

Revised version: These observations led us to hypothesize that phagocytosis, which heavily relies on actin rearrangement, might also be affected in Pik3caRBD/- BMDMs. To test this hypothesis, we assessed the ability of Pik3caRBD/- and Pik3caWT/- BMDMs to phagocytose fluorescent microspheres, Borrelia burgdorferi, and apoptotic cells. The results revealed varying outcomes depending on the target. When phagocytosing 1 μm non-opsonized green-fluorescent beads, Pik3caRBD/- BMDMs were less efficient at engulfing the microspheres compared to their WT counterparts (Figure 4-Supplement 5A). Similar results were obtained when Pik3caWT/WT BMDMs were treated with BYL-719 (Figure 4-Supplement 5A). In contrast, no significant differences were observed between Pik3caRBD/- and Pik3caWT/- BMDMs in their ability to phagocytose Borrelia burgdorferi (Figure 4-Supplement 5B).

Lastly, we evaluated the phagocytosis of apoptotic cells by Pik3caRBD/- BMDMs over extended periods. There were no differences in the initial uptake of apoptotic cells between Pik3caRBD/-and Pik3caWT/- BMDMs. However, at later time points, Pik3caRBD/- BMDMs showed an accumulation of phagocytosed particles. Similar results were obtained when Pik3caWT/WT BMDMs were treated with BYL-719. These data suggest a delay in processing and degradation, indicating a potential role for RAS-p110α in the later stages of phagocytosis (Fig. 4H).

In summary, our findings suggest that loss of RAS-p110α interaction leads to increased actin polymerization, resulting in stiffer and less deformable cells, which impairs phagocytic efficiency for certain targets. Specifically, RAS-p110α is important for the effective phagocytosis of non-biological particles and may also play a role in the proper processing and degradation of phagocytosed apoptotic cells. The differential effects on various phagocytic targets highlight the complexity of RAS-p110α’s role in macrophage biology and underscore the importance of cytoskeletal flexibility in efficient phagocytosis.

Disruption of RAS-p110α activation impacts the secretome of macrophages

Earlier version: Given our previous observations, we hypothesized that the cytoskeletal alterations observed after disruption of RAS-p110α interaction, may significantly impact the secretory functions of macrophages. The accumulation of apoptotic material in Pik3caRBD/- BMDMs suggests a disruption in normal phagocytic processing, which could influence the release of cytokines and other inflammatory mediators. Thus, we analysed the secretome of Pik3caRBD/- and Pik3caWT/- BMDMs both under steady state condition and during phagocytosis. We utilized apoptotic LKR10 cells, a murine lung cancer cell line, as the substrate in our phagocytosis assay. LKR10 cells were exposed to cisplatin for 16 h and apoptosis was confirmed by Western Blot (Supplementary Fig. S5A). Our goal was to create a more physiologically relevant experimental setting that more closely mimics the complex nature of the inflammatory response. For the secretome analysis, macrophages were incubated with or without apoptotic cells for 16 h. Culture supernatants were collected, clarified, and subjected to label-free quantitative proteomics analysis.

Revised version: Given our previous observations, we hypothesized that the cytoskeletal alterations observed after disruption of RAS-p110a interaction, may significantly impact the secretory functions of macrophages. The accumulation of apoptotic material in Pik3caRBD/- BMDMs suggests a disruption in normal phagocytic processing, which could influence the release of cytokines and other inflammatory mediators. Thus, we analysed the secretome of Pik3caRBD/- and Pik3caWT/- BMDMs both under steady state condition and during phagocytosis. We utilized apoptotic LKR10 cells, a murine lung cancer cell line, as the substrate in our phagocytosis assay. LKR10 cells were exposed to cisplatin for 16 h and apoptosis was confirmed by Western Blot (Figure 5-Supplement 6A). Our goal was to create a more physiologically relevant experimental setting that more closely mimics the complex nature of the inflammatory response. For the secretome analysis, macrophages were incubated with or without apoptotic cells for 16 h. Culture supernatants were collected, clarified, and subjected to label-free quantitative proteomics analysis.

Earlier version: A total of 127 peptides corresponding to 105 proteins (Supplementary Table 1) showed differential expression between Pik3caRBD/- and Pik3caWT/- BMDM secretomes at steady state conditions. Additionally, 359 peptides corresponding to 210 proteins (Supplementary Table 2) were present a significantly different levels in the secretomes of Pik3caRBD/- and Pik3caWT/- BMDM during phagocytosis of apoptotic cells. Next, we compared these peptides with the list of secreted proteins available at The Human Protein Atlas, and removed those that do not correspond to secreted proteins. After this step, 18 proteins were found to be differentially secreted by Pik3caRBD/- and Pik3caWT/- BMDM in steady state conditions (Fig. 5A), and 38 by Pik3caRBD/- and Pik3caWT/- BMDM during phagocytosis (Fig. 5B and Supplementary Fig. S5B and S5C). Surprisingly, most proteins were secreted at lower levels in Pik3caRBD/- BMDMs, independently of the conditions under study. 12 proteins were differentially secreted in both experimental conditions under study (Fig. 5C).

Revised version: A total of 127 peptides corresponding to 105 proteins (Supplementary File 1) showed differential expression between Pik3caRBD/- and Pik3caWT/- BMDM secretomes at steady state conditions. Additionally, 359 peptides corresponding to 210 proteins (Supplementary File 2) were present a significantly different levels in the secretomes of Pik3caRBD/- and Pik3caWT/- BMDM during phagocytosis of apoptotic cells. Next, we compared these peptides with the list of secreted proteins available at The Human Protein Atlas, and removed those that do not correspond to secreted proteins. After this step, 18 proteins were found to be differentially secreted by Pik3caRBD/- and Pik3caWT/- BMDM in steady state conditions (Fig. 5A), and 38 by Pik3caRBD/- and Pik3caWT/- BMDM during phagocytosis (Fig. 5B and Figure 5-Supplement 6B and 6C). Surprisingly, most proteins were secreted at lower levels in Pik3caRBD/- BMDMs, independently of the conditions under study. 12 proteins were differentially secreted in both experimental conditions under study (Fig. 5C).

Secretome analysis of Pik3caRBD/- BMDMs suggested a defect in complement activation and lysosomal function.

Earlier version: (A and B) Volcano plots of secretome analysis from Pik3caWT/- and Pik3caRBD/- BMDMs in nonstimulated conditions (A) or phagocytosing apoptotic cells (B). The x-axis shows the log2 FC of each identified protein and the y-axis the corresponding -log10 P value. Statistically significant peptides with FC ≥2 and p-value < 0.05 are in blue; peptides that do not pass this threshold are in grey; peptides with FC ≥ 2 but p-value ≥ 0.05 are in green; (C) Venn diagram showing the overlap of Pik3caRBD/- vs Pik3caWT/- BMDMs differentially expressed proteins at rest and during phagocytosis of apoptotic cells. Proteins displayed in the blue square represents peptides differentially expressed in resting BMDMs, the brown square represents peptides differentially expressed during phagocytosis, and the peptides differentially expressed in both conditions are displayed in purple; (D and E) Network analysis of significantly expressed proteins identified in the secretome analysis of Pik3caRBD/- vs Pik3caWT/- BMDMs in steady state conditions (D) or during phagocytosis of apoptotic cells (E). The nodes represent individual proteins, and the edges represent known interactions between proteins, either coexpression (purple) or physical interaction (thick pink). The size of each node reflects the significance of differential expression. Proteins in green have been implicated in lysosomal function, while proteins in orange are members of the complement cascade.

Revised version: (A and B) Volcano plots of secretóme analysis from Pik3caWT/- and Pik3caRBD/- BMDMs in nonstimulated conditions (A) or phagocytosing apoptotic cells (B). The x-axis shows the log2 FC of each identified protein and the y-axis the corresponding -log10 P value. Statistically significant peptides with FC ≥2 and p-value < 0.05 are in blue; peptides that do not pass this threshold are in grey; peptides with FC ≥ 2 but p-value ≥ 0.05 are in green; (C) Venn diagram showing the overlap of Pik3caRBD/- vs Pik3caWT/- BMDMs differentially expressed proteins at rest and during phagocytosis of apoptotic cells. Proteins displayed in the blue square represents peptides differentially expressed in resting BMDMs, the brown square represents peptides differentially expressed during phagocytosis, and the peptides differentially expressed in both conditions are displayed in purple; (D and E) Network analysis of significantly expressed proteins identified in the secretome analysis of Pik3caRBD/- vs Pik3caWT/- BMDMs in steady state conditions (D) or during phagocytosis of apoptotic cells (E). The nodes represent individual proteins, and the edges represent known interactions between proteins, either coexpression (purple) or physical interaction (thick pink). The size of each node reflects the significance of differential expression. Proteins in green have been implicated in lysosomal function, while proteins in orange are members of the complement cascade.

Functional analysis of differentially secreted proteins in unstimulated BMDMs showed no significant pathways related to these group of proteins (Fig. 5D). However, differentially secreted proteins in phagocytosing conditions were involved in two main biological processes: complement activation (C1qa, C1qb and C1qc connected through physical interaction and C3, C9 and Cfb through coexpression) and lysosome function, mainly cathepsins (Ctsd, Ctsb, Ctsz, Cst3, Psap, Anxa1 and Gsn) (Fig. 5E). Together, the complement cascade and lysosome function work in concert to provide an effective defence against pathogens and promote overall maintenance of cellular homeostasis22, 45–47 and data from the secretome analysis of Pik3caRBD/- and Pik3caWT/- BMDMs suggests that RAS activation of p110α may play a crucial role in the regulation of both response pathways. Moreover, these findings align with previous observations showing an accumulation of phagocytosed material in Pik3caRBD/- BMDMs, indicating a potential disruption in the processing and degradation of engulfed targets.

Disruption of RAS-p110α signalling leads to altered lysosomal function

We next aimed to functionally validate the proteomics data suggesting that RAS-p110α activation regulates lysosomal function in macrophages. First, we performed immunofluorescence analysis of lysosomal-associated membrane protein 1 (LAMP1) to assess lysosomal biogenesis and function in Pik3caRBD/- and Pik3caWT/- BMDMs. Analysis of LAMP1 in Pik3caRBD/- and Pik3caWT/- BMDMs showed a decrease in LAMP1 expression in steady state conditions (Fig. 6A) suggesting a decrease in the number of lysosomes after disruption of RAS-p110α interaction. However, Pik3caRBD/- showed increased levels of Lamp1 expression when subjected to phagocytosis of apoptotic cells, suggesting an increase in the number of phagolysosomes present in these cells.

Earlier version: Disruption of RAS-p110α activation in BMDMs leads to abnormal lysosomal function

Revised version: Disruption of RAS-p110a activation in BMDMs leads to abnormal lysosomal function

(A) Representative IF images and quantification analysis of lamp intensity in Pik3caWT/- and Pik3caRBD/- BMDMs in steady state conditions and during phagocytosis of apoptotic LKR10 cells. Three independent biological replicates were analysed (n ≥ 250 total cells). Error bars indicate mean ± SEM; (B) Quantification analysis of lysosomal function in Pik3caWT/-, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 BMDMs unstimulated or activated with LPS + IFN-γ using Lysotracker staining. Three independent biological replicates were analysed (n ≥ 250 total cells); (C) Violin plot displaying lysosome pH acidity in unstimulated and LPS+IFN-γ stimulated Pik3caWT/- (black), Pik3caRBD/- (red) and Pik3caWT/WT BMDMs treated with BYL719 (Blue) BMDMs, as determined by Lysosensor staining. Three independent biological replicates were analysed (n ≥ 250 total cells); (D) Pik3caWT/- and Pik3caRBD/- BMDMs were activated with LPS and IFN-γ and expression and activation of Cathepsin B and D were blotted; (E) Quantification of phagocytosed apoptotic cells in Pik3caWT/-, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 BMDMs. Apoptotic cells were labelled with a red tracker and allowed to be phagocytosed by control, Pik3caRBD/- and Pik3caWT/WT BMDMs treated with BYL719 for 16 hours, measuring the amount of cell tracker that was internalized at different time points. Each dot of the graph indicates a different independent experiment. Error bars indicate mean ± SEM; (F) Representative image and quantification of Cathepsin D expression in the inflammatory abscess of Pik3caWT/- and Pik3caRBD/- paws injected with zymosan; (G) Representative image and quantification of Cathepsin D expression in the inflammatory abscess of paws from control and Pik3caWT/WT mice treated with BYL719 injected with zymosan.

Statistical significance was obtained using Mann-Whitney test: *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001.

We hypothesised that the increase in Lamp1 expression in BMDMs lacking RAS-p110α interaction during phagocytosis could be attributed to aberrant lysosomal function and phagolysosome retention. Thus, we next evaluated lysosomal activity by using the lysosomotropic dye lysotracker red coupled with epifluorescence analysis, since it is specifically taken up by acidic organelles. As such, its accumulation is proportional to the number of acidic vesicles. Data analysis showed that disruption of RAS-PI3K in BMDMs leads to a significant decrease in lysotracker uptake, both in unstimulated conditions and also after activation with LPS+IFN-γ (Fig. 6B). Similar results were obtained with control BMDMs treated with BYL719, the p110α specific inhibitor. Data analysis showed a decrease in the uptake of lysotracker after p110α inhibition (Fig. 6B), further suggesting that loss of p110α function in BMDMs results in altered lysosomal pH. To confirm that lysosomal pH of Pik3caRBD/- BMDMs was less acidic, we next stained Pik3caRBD/- and Pik3caWT/- BMDMs with Green lysosensor, a pH-sensitive dye that exhibit a pH-dependent increase in fluorescent intensity upon lysosomal acidification. As shown in Fig. 6C, Pik3caRBD/- BMDMs presented attenuated fluorescent intensity both in unstimulated and activated when compared with that in the control group, indicating that their lysosomal content is less acidic than in Pik3caWT/- BMDMs.

Earlier version: Considering that Pik3caRBD/- lysosomes were less acidic, we next investigated the expression level and activation of some of the cathepsins identified in the secretome analysis (Fig. 5) by western blotting. Cathepsins play a critical role in lysosomal protein degradation. They are initially synthesized as inactive precursors and are activated through proteolysis in the lysosome at a low pH48. We found a reduction in the expression and activation of cathepsin D and cathepsin B in Pik3caRBD/- BMDMs upon stimulation with LPS+IFN-γ (Figure 6D). This observation suggests that the impaired lysosomal pH observed in Pik3caRBD/- BMDMs could potentially account for the observed decrease in cathepsin activity.

Revised version: Considering that Pik3caRBD/- lysosomes were less acidic, we next investigated the expression level and activation of some of the cathepsins identified in the secretome analysis (Fig. 5) by western blotting. Cathepsins play a critical role in lysosomal protein degradation. They are initially synthesized as inactive precursors and are activated through proteolysis in the lysosome at a low pH48. We found a reduction in the expression and activation of cathepsin D and cathepsin B in Pik3caRBD/- BMDMs upon stimulation with LPS+IFN-γ (Fig. 6D). This observation suggests that the impaired lysosomal pH observed in Pik3caRBD/- BMDMs could potentially account for the observed decrease in cathepsin activity.

Acidification of lysosomes and cathepsin activation are critical steps for activation of the resolutive stage of the inflammatory response, so we next evaluated the ability of Pik3caRBD/- BMDM’s lysosomal compartment to degrade internalised particles. We set up another phagocytosis assay in which Pik3caRBD/- and Pik3caWT/- BMDMs would phagocytose apoptotic cells that were transduced with GFP (which is pH sensitive49) and also labelled with Celltracker Red CMTPX (pH insensitive). Pik3caRBD/- and Pik3caWT/- BMDMs were allowed to engulf these apoptotic cells for 16 h and after this time, apoptotic cells were eliminated by washes and BMDMs were analysed by flow cytometry at different time points to measure GFP signal. Our results showed that GFP signal is lost significantly faster in control BMDMs than in Pik3caRBD/- BMDMs (Fig. 6E). As expected, we did not observe any decay in the red tracker signal. Collectively, these data evidence that Pik3caRBD/- BMDMs exhibit a significant impairment in removing engulfed particles, that can be attributed to the absence of lysosome acidification.

We have shown a delayed clearance of apoptotic cells in Pik3caRBD/- mice after zymosan injection (Fig. 1). This, together with previous data took us to analyse the levels of Cathepsin D in the inflamed abscess from the paws of Pik3caRBD/- and Pik3caWT/- mice. Results showed a significant decrease in the levels of Cathepsin D in the inflamed area of Pik3caRBD/- mice (Fig. 6F) and BYL719-treated mice (Fig. 6G).

Earlier version: Lysosomal digestion plays a pivotal role in activation of resolutive programs by mediating PPAR activation 50, 51. Analysis of PPARδ and PPARγ expression in phagocyting macrophages showed a significant decrease in the expression of PPARδ, but no PPARγ in Pik3caRBD/- BMDMs (Supplementary Fig. S6A), suggesting that resolutive programs might not be activated effectively.

Revised version: Lysosomal digestion plays a pivotal role in activation of resolutive programs by mediating PPAR activation 50, 51. Analysis of PPARδ and PPARγ expression in phagocyting macrophages showed a significant decrease in the expression of PPARδ, but no PPARγ in Pik3caRBD/- BMDMs (Figure 6- Supplement 7A), suggesting that resolutive programs might not be activated effectively.

In summary, our findings underscore the crucial role of RAS-p110α signalling axis in maintaining the balance of the inflammatory response and promoting timely resolution, providing further evidence for the significant involvement of RAS-p110α signalling in the response to inflammatory stimuli.

Discussion

Added in the revised version: In this study, we provide compelling evidence that disruption of RAS-p110α signalling or chemical inhibition of p110α impairs response to inflammatory stresses due to defects in both monocyte extravasation during the early stages of the inflammatory response and decreased lysosomal function during the later stages (Fig. 7). Although RAS-p110α signaling disruption affects all the hematopoietic cells, several lines of evidence indicate that macrophages play a central role in the observed phenotype. In our chimera mice experiments, neutrophil depletion did not alter the reduced transendothelial extravasation, suggesting that macrophages are the primary cell type involved. Furthermore, in the paw edema model, macrophages were the population exhibiting the most significant functional defects. Together, these findings, along with the specific deficiencies observed in myeloid populations, support a predominant role of macrophages in mediating the impaired inflammatory response. Nevertheless, we acknowledge that other myeloid cells, such as dendritic cells or additional immune populations, may also contribute to the global phenotype and further analysis should be performed to address this.

Added in the revised version: The role of RAS-p110α in macrophage-mediated inflammatory responses.

Added in the revised version: In the presence of functional RAS-p110α, monocytes efficiently migrate through the endothelium in response to inflammatory signals. Following extravasation, monocyte-derived and resident macrophages clear apoptotic cells via phagocytosis and lysosomal degradation, facilitating resolution of inflammation. In contrast, loss of RAS-p110α leads to cytoskeletal changes that hinder monocyte transendothelial migration. Additionally, resident macrophages exhibit impaired degradation of phagocytosed particles, resulting in unresolved inflammation and more severe acute inflammatory responses.

Removed in the revised version: In this study, we provide compelling evidence that disruption of RAS-p110α signalling or chemical inhibition of p110α impairs response to inflammatory stresses due to defects in both monocyte extravasation during the early stages of the inflammatory response and decreased lysosomal function during the later stages.

Earlier version: Monocyte extravasation during the inflammatory response is a critical step that allows immune cells to reach the site of infection or injury19. Impairment on this process causes slower or inadequate immune response, as macrophages are crucial for detecting and engulfing pathogens or cellular debris at the site of inflammation, as well as delayed resolution of inflammation, resulting in prolonged inflammation and potential tissue damage 52, 53. Our data shows that RAS binding to p110α is involved in macrophage extravasation by modulating actin dynamics. During monocyte extravasation, actin filaments form actin-rich protrusions that are essential for monocyte migration across the endothelium41, 54. Our data show that, during extravasation, RAS-p110α signalling regulates actin dynamics so monocytes are able to squeeze through endothelial cells41, 43.

Revised version: Monocyte extravasation during the inflammatory response is a critical step that allows immune cells to reach the site of infection or injury19. Impairment on this process causes slower or inadequate immune response, as macrophages are crucial for detecting and engulfing pathogens or cellular debris at the site of inflammation, as well as delayed resolution of inflammation, resulting in prolonged inflammation and subsequent tissue damage 52, 53.

Added in the revised version: Our data shows that RAS binding to p110α is involved in macrophage extravasation by modulating actin dynamics. During monocyte extravasation, actin filaments form actin-rich protrusions that are essential for monocyte migration across the endothelium41, 54. Our data show that, during extravasation, RAS-p110α signalling regulates actin dynamics so monocytes are able to squeeze through endothelial cells41, 43.

The disruption of RAS-p110α in macrophages emerges as a pivotal determinant in cellular mechanics, elucidating a cascade of events resulting in increased F-actin levels. The increase in cytoskeletal components, particularly F-actin, causes a notable increase in cell stiffness and a concurrent decrease in cell deformability. The orchestration of these changes underscores the intricate balance in cytoskeletal dynamics. Notably, the regulatory role of Rho-GTPases comes into focus as potential mediators of this phenomenon. Rho-GTPases are well-known mediators of actin cytoskeletal rearrangements55, and their dysregulation is known to impact cell mechanics56, 57. The observed increase in cell stiffness and deformability, may thus be governed by the modulation of Rho-GTPase activity. This intricate interplay provides a compelling avenue for further investigation into the molecular mechanisms through which RAS-p110α disruption influences Rho-GTPases, ultimately shaping the biomechanical properties of macrophages and may open avenues for therapeutic interventions targeting cytoskeletal dynamics in macrophages.

Monocyte extravasation shares many features with monocyte egression from the bone marrow. During monocyte egress, monocytes also undergo changes in cytoskeleton dynamics to detach from the sinusoidal endothelial cells and to migrate through the endothelial fenestrae and basement membrane into the bone marrow sinusoids58. Thus, the decrease in the number of classical monocytes in the Pik3caRBD/- mice in blood may be due, at least in part, to a defect in the extravasation process. This result may also explain, at least partially, the lack of macrophage recruitment to lung tumours found in previous studies using the RBD mouse model11, 12.

Our findings also revealed a crucial role of RAS-p110α activation in the acute phase of the inflammatory response by regulating effective lysosomal degradation of phagocytosed and engulfed material. Phagocytosis constitutes a vital step in the inflammatory response, whereby phagosomes bind to lysosomes to form phagolysosomes in which pathogens are eradicated to facilitate an appropriate host response. This response encompasses antigen presentation to engage T-cell responses, secretion of inflammatory mediators that guide the adaptive immune response, and initiation of tissue repair mechanisms 22, 47. Our data provide evidence that the activation of RAS-p110α signalling pathway is involved in the critical process of lysosomal acidification, which is essential for the efficient degradation of internalized particles and the activation of proteolytic enzymes, ultimately resulting in the formation of fully functional lysosomes. Consequently, lack of lysosome acidification impairs the expression and activation of important proteases such as cathepsin B and cathepsin D. The proper functioning of lysosomes is essential for mounting a robust response to inflammatory stress. Lysosomes play a central role in the breakdown of pathogens and dead cells by providing the necessary degradative enzymes and maintaining an acidic environment that facilitates the degradation of engulfed particles20, 21. When lysosomal function is compromised in macrophages, the degradation of phagocytosed material becomes impaired, leading to the accumulation of toxic debris. This accumulation subsequently triggers inflammation and causes damage to the surrounding tissues, leading to chronic inflammation and appearance of lysosomal storage diseases (LSD)59, 60.

Earlier version: Our data showed that the disruption of Ras-p110α in macrophages presents a multifaceted impact in inflammatory response, notably manifesting as a reduction in NF-κB activation and cathepsin expression coupled with diminished lysosomal function. NF-κB, a pivotal regulator of inflammatory responses61, 62, is intricately linked to lysosomal dynamics, with lysosomes playing a crucial role in modulating NF-κB signalling through the degradation of relevant molecules63, 64. Additionally, the observed decrease in NF-κB activation aligns with the decline in lysosomal function, which has previously been suggested in previous reports65, 66. This bidirectional influence, where perturbations in lysosomal function coincide with alterations in NF-κB activity, underscores the intricate and context-dependent nature of this regulatory network, suggesting a plausible mechanistic interdependence. Our data suggest that disruption in Ras-p110α signalling cascade could trigger downstream effects impacting both NF-κB activity and lysosomal integrity. This intricate nexus between Ras-p110α signalling, NF-κB activation, cathepsin expression, and lysosomal function underscores the complexity of cellular regulatory networks. Further analysis of the molecular pathways connecting these observations holds the potential to reveal novel insights into the coordinated regulation of inflammatory and lysosomal processes in macrophages. Thus, further investigations are necessary to decipher the causal relationships and unravel the broader implications of Ras-p110α disruption in shaping these intricate cellular responses during the inflammatory response.

Revised version: Our data showed that the disruption of Ras-p110α in macrophages presents a multifaceted impact in inflammatory response, notably manifesting as a reduction in NF-κB activation and cathepsin expression coupled with diminished lysosomal function. NF-κB, a pivotal regulator of inflammatory responses61, 62, is intricately linked to lysosomal dynamics, with lysosomes playing a crucial role in modulating NF-κB signalling through the degradation of relevant molecules63, 64. Additionally, the observed decrease in NF-κB activation aligns with the decline in lysosomal function, which has previously been suggested in previous reports65, 66. This bidirectional influence, where perturbations in lysosomal function coincide with alterations in NF-κB activity, underscores the intricate and contextdependent nature of this regulatory network, suggesting a plausible mechanistic interdependence. Our data suggest that disruption in Ras-p110α signalling cascade could trigger downstream effects impacting both NF-κB activity and lysosomal integrity. This intricate nexus between Ras-p110α signalling, NF-κB activation, cathepsin expression, and lysosomal function underscores the complexity of cellular regulatory networks. Further analysis of the molecular pathways connecting these observations holds the potential to reveal novel insights into the coordinated regulation of inflammatory and lysosomal processes in macrophages. Thus, further investigations are necessary to decipher the causal relationships and unravel the broader implications of Ras-p110α disruption in shaping these intricate cellular responses during the inflammatory response.

The possible functional link between the increased actin polymerization and lysosomal dysfunction observed in Pik3caRBD/- mice remains an unanswered question. Lysosome acidification and actin polymerization are tightly interconnected processes that have pivotal roles in numerous cellular functions67–69. Studies have demonstrated that lysosome acidification can influence actin polymerization dynamics through the activation of specific actin-binding proteins and the modulation of actin-regulatory proteins67, 70. Conversely, disruption in lysosome acidification, such as impaired proton pump activity or lysosomal storage disorders, have been associated with changes in actin polymerization and the organization of the cytoskeleton. Notably, it has been reported that depolimeryzation of F-actin plays a crucial role in assembling the macromolecular components of the acidification machinery in nascent endosomes71. Therefore, the intricate relationship between lysosome acidification and actin polymerization suggests a potential reciprocal influence, where perturbations in one process could impact the other. Further investigation is required to determine the precise regulatory mechanisms by which p110α influences both lysosome acidification and actin polymerization, whether it occurs in a linear manner or through separate pathways.

This study provides valuable insights into the mechanisms that govern the immune response to inflammation, particularly emphasizing the essential role of RAS-p110α signaling. Our results underscore the pivotal role of RAS-p110α signaling in both the initiation and resolution of inflammation. The impaired ability of RAS-PI3K-deficient monocytes and macrophages to effectively migrate, initiate a robust inflammatory response, and resolve inflammation highlights the critical importance of this pathway in maintaining immune homeostasis. The persistent inflammation observed in these models may contribute to the development and perpetuation of chronic inflammatory conditions. Given the centrality of p110α in both initiating and resolving inflammation, its dysfunction could be a key factor in the chronic, dysregulated inflammation characteristic of diseases such as rheumatoid arthritis, inflammatory bowel disease, psoriasis, or systemic lupus erythematosus.

The potential of p110α as a therapeutic target in these conditions is especially intriguing. The recent identification of a p110α small molecule activator72 offers a promising tool to explore this avenue. By transiently enhancing p110α function, it may be possible to promote the resolution of inflammation and thereby alleviate the symptoms and progression of these inflammatory diseases. Further research into this therapeutic strategy could open new pathways for treating chronic inflammatory disorders, providing much-needed relief for patients affected by these debilitating conditions.